---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: a0189e806d
xp: 50
```

`@projector_key`
32e84a16b89080fc0de303d0a6b1521a

---

## Bag-of-words

```yaml
type: VideoExercise
key: b0bb07659c
xp: 50
```

`@projector_key`
11bedc6349d6ea75bdb16ef8a9289b1f
